## Before submitting
1. Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

2. Make sure that no assertions fail or exceptions occur, otherwise points will be subtracted.

3. After you submit the notebook more tests will be run on your code. The fact that no assertions fail on your computer localy does not guarantee that you completed the exercise correctly.

4. After a function has been tested you can assume that for the following cells a correct implementation is used.

5. Please submit only the edited original `*.ipynb` file.

6. Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE". Edit only between `YOUR CODE HERE` and `END YOUR CODE`.

7. Make sure to use Python 3, not Python 2.

8. Only work on the exercises using Jupyter Notebook. While tools such as PyCharm and VSCode support the `ipynb` format, they overwrite crucial metadata, which break the autograder system.

9. Do **NOT** under any circustances delete any cells that you didn't insert yourselves. If you accidentally delete a cell either undo the deletion using `Edit/Undo Delete Cells` or redownload the notebook file from ISIS and paste your existing code in there. 


Fill your group name and members below:

```python
GROUPNAME = ""
COLLABORATORS = ""
```

# Exercise Sheet 3: Sampling, Simulation

In this exercise sheet, we will simulate a Markov chain. In the first part, we will consider a pure Python based implementation where a single particle jumps from one position to another of the lattice, where all transitions to neighboring states have the same probability. Then, we will add probabilities for the transitions. Finally, the implementation will be parallelized to run many chains in parallel.

Keep in mind that you can get the documentation of modules and functions using the `help()` function.

```python
%matplotlib inline
import random 

# For testing
import unittest; t = unittest.TestCase()
```

<!-- grade_id: cell-4d484b3190814b72, points: 0, grade: true, solution: false -->
```python
'''
This cell is for grading.DO NOT remove it
'''
```

## Exercise 1: Random moves in a lattice (20 P)

In this exercise, we will simulate the propagation of particles in a graph composed of 8 states (denoted by letters A-H) and stored in the variable `S` defined in the cell below. The lattice is the following:

![](lattice.png)

The particle starts in state `A` and then jumps randomly from its current state to one of its neighbors, all with same probability. Note that it cannot stay at the current position. The dictionary `T` defined in the cell below encode such transition behavior.

```python
# List of states
S = list('ABCDEFGH')

# Dictionary of transitions
T = {'A': 'BE', 'B': 'AFC', 'C': 'BGD', 'D': 'CH', 'E': 'AF', 'F': 'EBG', 'G': 'FCH', 'H': 'GD'}
print("States: {}".format(S))
print("Transitions: {}".format(T))
```

Using pure Python, set the initial state to `A` and run it for 1999 iterations. Return the sequence of states visited by the particle as a list. Set the random seed of the module `random` to value `123` using the function `random.seed` before starting the simulation in order to produce deterministic results.

<!-- grade_id: cell-be8591ff18527982, points: 0, grade: false, solution: true -->
```python
def simulate(transitions):
    '''
    Simulates a markov chain defined by the above transitions.
    This function always sets the random seed to `123`. All simulations start with 
    initial state `A`. It always simulates 2000 steps including the initial state.
    Args:
        transitions (dict): A dictionary with eight keys [A-H]. For each key a string is 
        mapped as its value. Each of those strings can only contain the letters [A-H] each
        letter can only appear once. `'A': 'BE'` means that from state `A` we can reach 
        the states `B` and `E` and no other state. 
    Returns:
        list: A list of states (a string containing one of the letters [A-H]) 
        that were visited during the simulation.
    '''
    # YOUR CODE HERE
    raise NotImplementedError("Relplace this line with your code")
    # YOUR CODE HERE
    
    
```

<!-- grade_id: cell-e189a7a547cbdcbc, points: 5, grade: true, solution: false -->
```python
# Tests
X = simulate(T)

# Print the first 10 states
print(X[:10])
t.assertIsInstance(X, list, 'The state sequence must be a list')
t.assertIsInstance(X[0], str, 'The state sequence must only contain strings')
t.assertEqual(X[0], 'A', 'The state sequence must start with A')
t.assertTrue(all(x in 'ABCDEFGH' for x in X), 'Your state sequence containes an invalid state')
t.assertEqual(len(X), 2000)
```

<!-- grade_id: cell-5749c0c7ceef3507, points: 0, grade: true, solution: false -->
```python
'''This cell is for grading. DO NOT remove it'''
```

Implement a function that returns a list of the relative frequencies of each state.

<!-- grade_id: cell-d2abc424da288db8, points: 0, grade: false, solution: true -->
```python
import numpy as np
def compute_histogram(state_sequence):
    '''
    Returns a list of percentages relating as to how many times each state has been visited
    according to the `state_sequence` list
    
    Args:
        state_sequence (list): A list of states
    Returns:
        list: A list of length 8. Contains the percentage `[0-1]` of occurances of each state
        in the `state_sequence`. 
    '''
    # YOUR CODE HERE
    raise NotImplementedError("Relplace this line with your code")
    # YOUR CODE HERE
    
```

<!-- grade_id: cell-a083af994b2e82c2, points: 10, grade: true, solution: false -->
```python
h = compute_histogram(X)
t.assertIsInstance(h, list)
# Check if the histogram is a valid probability distribution
np.testing.assert_allclose(np.sum(h), 1.)

```

<!-- grade_id: cell-5c1b03dc60626965, points: 0, grade: true, solution: false -->
```python
'''This cell is for grading.DO NOT remove it'''
```

Using the above `compute_histogram` function, produce a bar plot using `matplotlib` (`matplotlib.pyplot.bar`) showing the fraction of the time the particle is found in a given state, averaged over the whole simulation. Do **not** call plt.show in the function (it is automatically called because of `%matplotlib inline`

<!-- grade_id: cell-4e935488f89005fd, points: 5, grade: true, solution: true -->
```python
import matplotlib.pyplot as plt
def plot_histogram(histogram):
    '''
    Plots a bar graph of a provided histogram. 
    
    Args:
        histogram: A list of propabilites (values in `[0-1]`)
    Returns: 
        None
    '''
    # YOUR CODE HERE
    raise NotImplementedError("Relplace this line with your code")
    # YOUR CODE HERE
    
```

<!-- grade_id: cell-8adc44a487ec126c, points: 0, grade: true, solution: false -->
```python
# Plot the histogram of the above defined sequence X
h = compute_histogram(X)
plot_histogram(h)
```

## Exercise 2: Adding a special state (20 P)

Suppose now that the rule (defined by the transition dictionary) is modified such that everytime the particle is in state `F`, it always moves to `E` in the next step.

* Modify the code to handle this special case, and create a bar plot for the new states distribution. Make sure to not modify the original transition dictionary. To achieve this, you can have a look at the `copy` module, but every working solution is accepted.

<!-- grade_id: cell-b23ac9c3c7d18076, points: 0, grade: false, solution: true -->
```python
def modify_transitions(transitions):
    '''
    Creates a modified transition dictionary without modifing the provided one.
    
    This function creates a new transition dictionary such that from state `F` the only 
    possible following state is `E`.
    
    Args:
        transitions (dict): A dictionary that describes the possible transitions from each state
    Returns:
        dict: A modified transition dict where from state `F` only state `E` can follow
     '''
    # YOUR CODE HERE
    raise NotImplementedError("Relplace this line with your code")
    # YOUR CODE HERE
    
    
```

```python
new_T = modify_transitions(T)
new_X = simulate(new_T)
h = compute_histogram(new_X)
plot_histogram(h)
```

<!-- grade_id: cell-4c8eda5748d2ac39, points: 20, grade: true, solution: false -->
```python
# Tests
new_T = modify_transitions(T)
print('T[\'F\'] = {}'.format(T['F']))
print('new_T[\'F\'] = {}'.format(new_T['F']))
t.assertIsInstance(new_T, dict)
t.assertIsNot(T, new_T, 'T and new_T should not be the same instance')
```

## Exercise 3: Exact solution to the previous exercise (20 P)

For simple Markov chains, a number of statistics can be obtained analytically from the structure of the transition model, in particular, by analysis of the transition matrix.

* Compute the transition matrices associated to the models of exercise 1 and 2 (make sure that each row in these matrices sums to 1).
* Give the transition matrices as argument to the function `utils.getstationary(P)` and return their result.

This last function computes in closed form the stationary distribution associated to a given transition matrix `P` (i.e. the one we would get if running the simulation with such transition matrix for infinitely many time steps and looking at state frequencies).

<!-- grade_id: cell-e04eeaf399e5473f, points: 0, grade: false, solution: true -->
```python
def to_matrix(transition):
    '''
    Converts a transition dictionary into a transition matrix. The first row represents
    the propability of moving from the first state to every state.
    
    If the state dict is irreflexive (we cannot go from one state to the same state) 
    the sum of the diagonal is 0. 
    
    The sum of each row should be 1.
    
    All the elements in the matrix are values in [0-1].
    
    Args:
        transition (dict): A dictionary describing the possible transitions from each state.
        
    Returns:
        np.ndarray: The transition matrix (ndim=2) that represents the same (uniform) transitions as the 
        transition dict
    '''
    # YOUR CODE HERE
    raise NotImplementedError("Relplace this line with your code")
    # YOUR CODE HERE
    
    
```

<!-- grade_id: cell-c56cbb51974a1435, points: 15, grade: true, solution: false -->
```python
'''
Test `to_matrix`
Do not delete this cell or your submission will fail
'''
matrix_T = to_matrix(T)
print(np.round(matrix_T, 2))

t.assertIsInstance(matrix_T, np.ndarray)
np.testing.assert_allclose(np.sum(matrix_T, axis=1), 1.0)

```

<!-- grade_id: cell-b7872f409809f378, points: 0, grade: true, solution: false -->
```python
'''This cell is for grading. DO NOT remove it'''
```

<!-- grade_id: cell-16d4205430adc58c, points: 0, grade: false, solution: true -->
```python
import utils

def get_stationary(matrix):
    '''
    Wrapper function for `utils.getstationary`
    
    Args:
        matrix (np.ndarray): Transition matrix (ndim=2)
    Returns:
        np.ndarray: Stationary distribution vector (ndim=1) of `matrix`.
    '''
    # YOUR CODE HERE
    raise NotImplementedError("Relplace this line with your code")
    # YOUR CODE HERE
    
```

<!-- grade_id: cell-269b5f3f15a5f64d, points: 5, grade: true, solution: false -->
```python
s = get_stationary(matrix_T)
print(s.round(3))
np.testing.assert_allclose(s.sum(), 1)
```

<!-- grade_id: cell-c1ee593c8d0e538b, points: 0, grade: true, solution: false -->
```python
'''This cell is for grading. DO NOT remove it'''
```

## Exercise 4: Adding non-uniform transition probabilities (20 P)

We consider the original lattice defined by the variable `T`. We set transition probabilities for each state to be such that:

1. The probability of moving vertically is the same as the probability of moving horizontally for states that have only two connections.
2. For states with 3 connections the probability of moving vertically is 0.5 while the probability of moving left is always twice the probability of moving right.

**Hints**: You can notice that for each state, the transitions towards other states are always listed from left to right in the dictionary `T`. Also note that characters A-H can be mapped to integer values using the Python function `ord()`, thus, giving a direct relation between state names and indices of the transition matrix.

**a)** Build a function that converts the string state into a numeric index

<!-- grade_id: cell-ed06d04f24b9db06, points: 0, grade: false, solution: true -->
```python
def state_string_to_index(state):
    '''
    Converts the state string into a numerical index, where:
    'A' -> 0
    'B' -> 1
    ...
    'H' -> 7
    
    Args:
        state (string): A state string in [A-H] with len(state)=1
    Returns:
        int: The index of the state in [0-7]
    '''
    # YOUR CODE HERE
    raise NotImplementedError("Relplace this line with your code")
    # YOUR CODE HERE
    
```

<!-- grade_id: cell-d5a69e2ff33dde8a, points: 5, grade: true, solution: false -->
```python
# Test state_string_to_index
A_idx = state_string_to_index('A')
t.assertNotIsInstance(A_idx, float)
t.assertEqual(A_idx, 0)
```

<!-- grade_id: cell-f8a4376a239ee7e1, points: 0, grade: true, solution: false -->
```python
'''This cell is for grading. DO NOT remove it'''
```

**b**) Now implement the `build_transition_matrix` according to the rules defined above.

<!-- grade_id: cell-1560bb06a55de020, points: 0, grade: false, solution: true -->
```python
def build_transition_matrix(transition):
    '''
    Builds a transition matrix from a transition dictionary, similarly to
    `to_matrix` function. However, this function does not create a uniform distribution
    among the following states. 
    If the the next valid states are two then the distribution is uniform.
    If the the next valid states are three, then moving vertically should have a 
    50% chance and moving left twice as much as moving right. 
    
    Like in the `to_matrix` function the sum of each row should be 1.
    
    Args:
        transition (dict) A dictionary describing the possible transitions from each state.
    Returns:
        np.ndarray: A transition matrix
    
    '''
    # YOUR CODE HERE
    raise NotImplementedError("Relplace this line with your code")
    # YOUR CODE HERE
    

```

<!-- grade_id: cell-cceb43e3ed998bfe, points: 15, grade: true, solution: false -->
```python
# Test build_transition_matrix
P = build_transition_matrix(T)
print(P.round(2))

t.assertIsInstance(P, np.ndarray)
np.testing.assert_allclose(P.sum(axis=1), 1.0)
```

<!-- grade_id: cell-f8d9ce6430c2dbb4, points: 0, grade: true, solution: false -->
```python
'''This cell is for grading. DO NOT remove it'''
```

## Exercise 5: Simulation for multiple particles (20 P)

We let 1000 particles evolve simultaneously in the system described in Exercise 4. The initial state of these particles is pseudo-random and given by the function `utils.getinitialstate()`.

* Using the function `utils.mcstep()` that was introduced during the lecture, simulate this system for 500 time steps.
* Estimate the stationary distribution by looking at the distribution of these particles in state space after 500 time steps.

For reproducibility, give seed values to the function utils.mcstep corresponding to the current time step of the simulation (i.e. from 0 to 499).

<!-- grade_id: cell-2af098b516309a41, points: 0, grade: false, solution: true -->
```python
import utils

def simulate_1000(transition):
    '''
    Simulates 1000 particles for 500 time steps, in order to approximate
    the stationary distribution
    
    Args:
        transition: A transition dict, that will be converted into a transition matrix using the 
        `build_transition_matrix` function
    Returns:
        np.ndarray: The estimated stationary distribution vector (ndim=1)
    
    '''
    # YOUR CODE HERE
    raise NotImplementedError("Relplace this line with your code")
    # YOUR CODE HERE
    
```

<!-- grade_id: cell-621b6e15bfc77266, points: 20, grade: true, solution: false -->
```python
# Tests
stationary_distribution = simulate_1000(T)
print(stationary_distribution)
t.assertIsInstance(stationary_distribution, np.ndarray)
np.testing.assert_allclose(np.sum(stationary_distribution), 1) 
```